# Taller 00 — Ensayo del flujo de entrega

**Procesos Estocásticos · Área de Modelado y Simulación · Universidad EIA**

---

Este taller **no tiene peso en la nota**. Su único propósito es que ensayes el flujo completo de trabajo antes del Taller 01.

## Flujo de trabajo

1. Resuelve el ejercicio en este notebook, dentro de tu entorno en DataHub EIA.
2. Pulsa el botón **Validate** en la barra superior de JupyterLab. Verás si las pruebas visibles pasan o fallan.
3. Corrige y vuelve a validar cuantas veces necesites. La validación es local: no envía nada ni consume intentos.
4. Cuando estés conforme, descarga el archivo `.ipynb` a tu computador (clic derecho sobre el archivo → *Download*).
5. Renómbralo como `<tu_código>_taller00.ipynb` (ejemplo: `1234567_taller00.ipynb`).
6. Súbelo a la tarea correspondiente en Moodle.

## Advertencia importante

El botón **Validate** ejecuta únicamente las pruebas **visibles**. La calificación final incluye además pruebas **ocultas** que verifican tu respuesta con mayor detalle.

> Que la validación aparezca en verde **no garantiza** la nota máxima. Es una señal de que vas bien encaminado, no un certificado de respuesta correcta.

## Reglas

- Escribe tu solución únicamente donde se indica `# TU CÓDIGO AQUÍ`.
- No modifiques ni elimines las celdas de prueba.
- No cambies los nombres de las variables solicitadas.

```python
import numpy as np

np.set_printoptions(precision=4, suppress=True)
```

---

## Contexto

Una máquina puede encontrarse en dos estados: **operativa** (estado 0) o **en falla** (estado 1). Su comportamiento de un día al siguiente se describe mediante la matriz de transición

$$
P = \begin{pmatrix} 0.9 & 0.1 \\ 0.5 & 0.5 \end{pmatrix}
$$

donde $P_{ij}$ es la probabilidad de pasar del estado $i$ al estado $j$ en un día.

Hoy la máquina está **operativa**, de modo que la distribución inicial es $\pi_0 = (1, 0)$.

---

### Ejercicio 1 (1 punto)

Define la matriz de transición como un arreglo de NumPy llamado `P`, y la distribución inicial como un arreglo llamado `pi_0`.

<!-- grade_id: ejercicio1, points: 0, grade: false, solution: true -->
```python
# TU CÓDIGO AQUÍ — elimina esta línea y la siguiente
raise NotImplementedError()
```

<!-- grade_id: test_ejercicio1_visible, points: 0.5, grade: true, solution: false -->
```python
# PRUEBA VISIBLE — verifica la forma de los arreglos
assert P.shape == (2, 2), 'P debe ser una matriz de 2x2.'
assert pi_0.shape == (2,), 'pi_0 debe ser un vector de longitud 2.'
assert np.allclose(P.sum(axis=1), 1.0), 'Cada fila de P debe sumar 1.'
assert np.isclose(pi_0.sum(), 1.0), 'pi_0 debe sumar 1.'
print('Prueba visible superada.')
```

<!-- grade_id: test_ejercicio1_oculto, points: 0.5, grade: true, solution: false -->
```python

```

---

### Ejercicio 2 (1 punto)

Calcula la distribución de estados **después de un día**, es decir $\pi_1 = \pi_0 P$. Guarda el resultado en una variable llamada `pi_1`.

*Sugerencia:* en NumPy el producto matricial se escribe con el operador `@`.

<!-- grade_id: ejercicio2, points: 0, grade: false, solution: true -->
```python
# TU CÓDIGO AQUÍ — elimina esta línea y la siguiente
raise NotImplementedError()

print(pi_1)
```

<!-- grade_id: test_ejercicio2_visible, points: 0.5, grade: true, solution: false -->
```python
# PRUEBA VISIBLE — verifica que pi_1 sea una distribución de probabilidad válida
assert pi_1.shape == (2,), 'pi_1 debe ser un vector de longitud 2.'
assert np.isclose(pi_1.sum(), 1.0), 'pi_1 debe sumar 1.'
assert np.all(pi_1 >= 0), 'Las probabilidades no pueden ser negativas.'
print('Prueba visible superada.')
```

<!-- grade_id: test_ejercicio2_oculto, points: 0.5, grade: true, solution: false -->
```python

```

---

## Cierre

Si las pruebas visibles pasaron, has completado el ejercicio. Recuerda que existen pruebas ocultas adicionales.

**Antes de entregar:**

1. Ejecuta *Kernel → Restart Kernel and Run All Cells* para verificar que el notebook corre de principio a fin sin errores.
2. Descarga el archivo `.ipynb`.
3. Renómbralo como `<tu_código>_taller00.ipynb`.
4. Súbelo a Moodle.

Cualquier dificultad con el flujo, repórtala antes del Taller 01. Ese es el propósito de este ensayo.