# Taller 01 — Simulación de cadenas de Markov

**Curso:** Procesos Estocásticos · Ingeniería Industrial · Universidad EIA
**Semana:** 2 · **Bloque:** A · **Puntos totales:** 15 (equivalente al primer tercio de "Laboratorios nbgrader Bloque A", 15% de la nota)
**Resultado de aprendizaje evaluado:** RA1 — *Modelar sistemas estocásticos discretos aplicando principios de probabilidad e ingeniería.*

## Contexto del problema

Una empresa de manufactura clasifica el estado de una máquina crítica en tres categorías, revisadas al final de cada turno:

- **Estado 0 — Operativa**
- **Estado 1 — En mantenimiento preventivo**
- **Estado 2 — Dañada (mantenimiento correctivo)**

A partir de datos históricos, el área de mantenimiento estimó las siguientes probabilidades de transición de un turno al siguiente:

| Desde \ Hacia | Operativa | Mantenimiento | Dañada |
|---|---|---|---|
| **Operativa** | 0.90 | 0.07 | 0.03 |
| **Mantenimiento** | 0.60 | 0.35 | 0.05 |
| **Dañada** | 0.10 | 0.50 | 0.40 |

## Instrucciones de entrega

- Completa únicamente las celdas marcadas con `# YOUR CODE HERE`.
- No modifiques las celdas de prueba (`### TEST`) — son las que usa nbgrader para calificar automáticamente.
- La Pregunta 3 se califica manualmente: escribe una explicación completa, no solo una palabra.
- Antes de entregar, corre todo el notebook de principio a fin (`Kernel → Restart & Run All`) para verificar que no queden errores.

---
## Ejercicio 1 — Construir la matriz de transición (4 pts)

**Relación con RA1:** este ejercicio evalúa tu capacidad de traducir una descripción verbal de un sistema estocástico a su representación matemática formal — el primer paso de todo modelado con cadenas de Markov.

<!-- grade_id: ex1_matriz_transicion, points: 4, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: ex1_test, points: 4, grade: true, solution: false -->
```python
### TEST
assert P.shape == (3, 3), "La matriz debe ser de 3x3"
assert np.allclose(P.sum(axis=1), 1), "Cada fila debe sumar 1 (son probabilidades)"
assert np.isclose(P[0, 0], 0.90), "P[0,0] (Operativa -> Operativa) es incorrecto"
assert np.isclose(P[2, 1], 0.50), "P[2,1] (Dañada -> Mantenimiento) es incorrecto"
print("Ejercicio 1: todas las verificaciones pasaron.")

```

---
## Ejercicio 2 — Simular la trayectoria y estimar la distribución estacionaria (6 pts)

**Relación con RA1:** aquí conectas la teoría (Semana 1: proceso estocástico como familia de variables aleatorias) con la práctica computacional que ya dominas desde Simulación (generación de trayectorias aleatorias). La distribución estacionaria te dice, en el largo plazo, qué fracción del tiempo la máquina pasa en cada estado — información directamente accionable para el área de mantenimiento.

Simula **10,000 turnos** empezando en el Estado 0 (Operativa), y estima la distribución estacionaria empírica (la fracción de turnos que la máquina pasó en cada estado).

<!-- grade_id: ex2_simulacion, points: 6, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: ex2_test, points: 6, grade: true, solution: false -->
```python
### TEST
assert len(trayectoria) == 10000, "La trayectoria debe tener 10000 pasos"
assert trayectoria[0] == 0, "Debe iniciar en el Estado 0 (Operativa)"
assert set(trayectoria).issubset({0, 1, 2}), "Solo deben aparecer los estados 0, 1, 2"
# La distribución estacionaria analítica de esta cadena es aprox. [0.82, 0.13, 0.05]
assert abs(distribucion_empirica[0] - 0.82) < 0.03, "La fracción del Estado 0 se aleja demasiado del valor esperado"
print("Ejercicio 2: todas las verificaciones pasaron.")

```

---
## Pregunta 3 — Interpretación (5 pts, calificación manual)

**Relación con RA1 y con Diligencia (marco 4D):** esta pregunta evalúa si puedes sostener la verificación de tu resultado más allá de "el código corrió sin error" — es decir, si discutes críticamente qué significa el número que obtuviste, no solo que lo obtuviste.

Responde en la celda de abajo (mínimo 100 palabras):

1. ¿Qué significa, en términos del negocio, que la probabilidad estacionaria del Estado 2 (Dañada) sea aproximadamente 0.06?
2. Si la simulación hubiera empezado en el Estado 2 en vez del Estado 0, ¿esperarías que la distribución estacionaria cambiara? Justifica tu respuesta con lo que sabes de la teoría (pista: piensa en qué significa "estacionaria").

<!-- grade_id: ex3_respuesta, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Bitácora de uso de IA (Diligencia)

Antes de entregar, documenta brevemente cómo usaste (o no) asistencia de IA en este taller — esto alimenta la competencia de Diligencia del marco 4D del área.

*¿Qué delegaste a IA (si algo)? ¿Qué verificaste tú mismo antes de confiar en el resultado?*

`Tu respuesta aquí (2-3 líneas)`